# Python review: Values, variables, types, lists, and strings

These first few notebooks are a set of exercises designed to reinforce various aspects of Python programming.

**Study hint: Read the test code!** You'll notice that most of the exercises below have a place for you to code up your answer followed by a "test cell." That's a code cell that checks the output of your code to see whether it appears to produce correct results. You can often learn a lot by reading the test code. In fact, sometimes it gives you a hint about how to approach the problem. As such, we encourage you to try to read the test cells even if they seem cryptic, which is deliberate!

**Debugging tip: Read assertions.** The test cells often run an `assert` statement to see whether some condition that it thinks should be true is true. If an assertion fails, look at the condition being checked and use that as a guide to help you debug. For example, if an assertion reads, `assert a + b == 3`, and that fails, inspect the values and types of `a` and `b` to help determine why their sum does not equal 3.

**Exercise 0** (1 point). Run the code cell below. It should display the output string, `Hello, world!`.

<!-- grade_id: hello_world_test, points: 1, grade: true, solution: false -->
```python
print("Hello, world!")
```

Output:
```
Hello, world!
```

**Exercise 1** (`x_float_test`: 1 point). Create a variable named `x_float` whose numerical value is one (1) and whose type is *floating-point* (i.e., `float`).

<!-- grade_id: x_float, points: 0, grade: false, solution: true -->
```python
###
### YOUR CODE HERE
x_float =1.0
###

```

<!-- grade_id: x_float_test, points: 1, grade: true, solution: false -->
```python
# `x_float_test`: Test cell
assert x_float == 1, f"`x_float` has the wrong value ({x_float} rather than 1.0)"
assert type(x_float) is float, f"`type(x_float)` == {type(x_float)} rather than `float`"
print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 2** (`strcat_ba_test`: 1 point). Complete the following function, `strcat_ba(a, b)`, so that given two strings, `a` and `b`, it returns the concatenation of `b` followed by `a` (pay attention to the order in these instructions!).

<!-- grade_id: strcat_ba, points: 0, grade: false, solution: true -->
```python
def strcat_ba(a, b):
    assert type(a) is str, f"Input argument `a` has `type(a)` is {type(a)} rather than `str`"
    assert type(b) is str, f"Input argument `b` has `type(b)` is {type(b)} rather than `str`"
    return b+a
###
### YOUR CODE HERE
###

```

<!-- grade_id: strcat_ba_test, points: 1, grade: true, solution: false -->
```python
# `strcat_ba_test`: Test cell

# Workaround:  # Python 3.5.2 does not have `random.choices()` (available in 3.6+)
def random_letter():
    from random import choice
    return choice('abcdefghijklmnopqrstuvwxyz')

def random_string(n, fun=random_letter):
    return ''.join([str(fun()) for _ in range(n)])

a = random_string(5)
b = random_string(3)
c = strcat_ba(a, b)
print('strcat_ba("{}", "{}") == "{}"'.format(a, b, c))
assert len(c) == len(a) + len(b), "`c` has the wrong length: {len(c)} rather than {len(a)+len(b)}"
assert c[:len(b)] == b
assert c[-len(a):] == a
print("\n(Passed!)")
```

Output:
```
strcat_ba("dmdwl", "xec") == "xecdmdwl"

(Passed!)
```

**Exercise 3** (`strcat_list_test`: 2 points). Complete the following function, `strcat_list(L)`, which generalizes the previous function: given a *list* of strings, `L[:]`, returns the concatenation of the strings in reverse order. For example:

```python
    strcat_list(['abc', 'def', 'ghi']) == 'ghidefabc'
```

<!-- grade_id: strcat_list, points: 0, grade: false, solution: true -->
```python
def strcat_list(L):
    assert type(L) is list
    a=""
    for i in range(len(L)-1,-1,-1):
        a=a+L[i]
    return a
    ###
    ### YOUR CODE HERE
    ###
strcat_list(['abc', 'def', 'ghi'])
```

Output:
```
'ghidefabc'
```

<!-- grade_id: strcat_list_test, points: 2, grade: true, solution: false -->
```python
# `strcat_list_test`: Test cell
n = 3
nL = 6
L = [random_string(n) for _ in range(nL)]
Lc = strcat_list(L)

print('L == {}'.format(L))
print('strcat_list(L) == \'{}\''.format(Lc))
assert all([Lc[i*n:(i+1)*n] == L[nL-i-1] for i, x in zip(range(nL), L)])
print("\n(Passed!)")
```

Output:
```
L == ['dxh', 'moa', 'ldv', 'kjm', 'kiy', 'ssp']
strcat_list(L) == 'sspkiykjmldvmoadxh'

(Passed!)
```

**Exercise 4** (`floor_fraction_test`: 1 point). Suppose you are given two variables, `a` and `b`, whose values are the real numbers, $a \geq 0$ (non-negative) and $b > 0$ (positive). Complete the function, `floor_fraction(a, b)` so that it returns $\left\lfloor\frac{a}{b}\right\rfloor$, that is, the *floor* of $\frac{a}{b}$. The *type* of the returned value must be `int` (an integer).

<!-- grade_id: floor_fraction, points: 0, grade: false, solution: true -->
```python
def is_number(x):
    """Returns `True` if `x` is a number-like type, e.g., `int`, `float`, `Decimal()`, ..."""
    from numbers import Number
    return isinstance(x, Number)
    
def floor_fraction(a, b):
    assert is_number(a) and a >= 0
    assert is_number(b) and b > 0
    ###
    ### YOUR CODE HERE
    return int(a/b)
    ###

```

<!-- grade_id: floor_fraction_test, points: 1, grade: true, solution: false -->
```python
# `floor_fraction_test`: Test cell
from random import random
a = random()
b = random()
c = floor_fraction(a, b)

print('floor_fraction({}, {}) == floor({}) == {}'.format(a, b, a/b, c))
assert b*c <= a <= b*(c+1)
assert type(c) is int, f"type(c) == {type(c)} rather than `int`"
print('\n(Passed!)')
```

Output:
```
floor_fraction(0.09838372637287807, 0.48274553902271666) == floor(0.20380038430194256) == 0

(Passed!)
```

**Exercise 5** (`ceiling_fraction_test`: 1 point). Complete the function, `ceiling_fraction(a, b)`, which for any numeric inputs, `a` and `b`, corresponding to real numbers, $a \geq 0$ and $b > 0$, returns $\left\lceil\frac{a}{b}\right\rceil$, that is, the *ceiling* of $\frac{a}{b}$. The type of the returned value must be `int`.

<!-- grade_id: ceiling_fraction, points: 0, grade: false, solution: true -->
```python
import math
def ceiling_fraction(a, b):
    assert is_number(a) and a >= 0
    assert is_number(b) and b > 0
    ###
    return math.ceil(a/b)
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: ceiling_fraction_test, points: 1, grade: true, solution: false -->
```python
# `ceiling_fraction_test`: Test cell
from random import random
a = random()
b = random()
c = ceiling_fraction(a, b)
print('ceiling_fraction({}, {}) == ceiling({}) == {}'.format(a, b, a/b, c))
assert b*(c-1) <= a <= b*c
assert type(c) is int
print("\n(Passed!)")
```

Output:
```
ceiling_fraction(0.6320891558989666, 0.982169433932816) == ceiling(0.6435642711542618) == 1

(Passed!)
```

**Exercise 6** (`report_exam_avg_test`: 1 point). Let `a`, `b`, and `c` represent three exam scores as numerical values. Complete the function, `report_exam_avg(a, b, c)` so that it computes the average score (equally weighted) and returns the string, `'Your average score: XX'`, where `XX` is the average rounded to one decimal place. For example:

```python
    report_exam_avg(100, 95, 80) == 'Your average score: 91.7'
```

<!-- grade_id: cell-a117d6495f42b850, points: 0, grade: false, solution: true -->
```python
import numpy as np
def report_exam_avg(a, b, c):
    #assert is_number(a) and is_number(b) and is_number(c)
    ###
    return 'Your average score: '+str(round((a+b+c)/3,1))
    ### YOUR CODE HERE
    ###
report_exam_avg(100, 95, 80)
```

Output:
```
'Your average score: 91.7'
```

<!-- grade_id: report_exam_avg_test, points: 1, grade: true, solution: false -->
```python
# `report_exam_avg_test`: Test cell
msg = report_exam_avg(100, 95, 80)
print(msg)
assert msg == 'Your average score: 91.7'

print("Checking some additional randomly generated cases:")
for _ in range(10):
    ex1 = random() * 100
    ex2 = random() * 100
    ex3 = random() * 100
    msg = report_exam_avg(ex1, ex2, ex3)
    ex_rounded_avg = float(msg.split()[-1])
    abs_err = abs(ex_rounded_avg*3 - (ex1 + ex2 + ex3)) / 3
    print("{}, {}, {} -> '{}' [{}]".format(ex1, ex2, ex3, msg, abs_err))
    assert abs_err <= 0.05

print("\n(Passed!)")
```

Output:
```
Your average score: 91.7
Checking some additional randomly generated cases:
98.15278963373629, 93.99490484408092, 20.488437745433586 -> 'Your average score: 70.9' [0.021289258916406577]
0.3854119769835296, 90.10129639730197, 2.340532940416873 -> 'Your average score: 30.9' [0.04241377156746277]
9.51796682006667, 21.917905308758556, 47.81445043029575 -> 'Your average score: 26.4' [0.01677418637366183]
38.3679068497781, 21.384165807208998, 81.40430602830776 -> 'Your average score: 47.1' [0.047873771568390566]
89.66825420299045, 69.74144493283002, 4.831189128112712 -> 'Your average score: 54.7' [0.04696275464438789]
45.731191172050224, 42.75255947073534, 39.21664332238633 -> 'Your average score: 42.6' [0.03320201160937112]
27.049266573266806, 94.61655951627725, 22.522337745300312 -> 'Your average score: 48.1' [0.03727872171854566]
32.77253989527969, 53.06070645018318, 18.531993433418215 -> 'Your average score: 34.8' [0.011586740372967105]
78.33242462903354, 60.495442856698745, 28.760391175
```

**Exercise 7** (`count_word_lengths_test`: 2 points). Write a function `count_word_lengths(s)` that, given a string consisting of words separated by spaces, returns a list containing the length of each word. Words will consist of lowercase alphabetic characters, and they may be separated by multiple consecutive spaces. If a string is empty or has no spaces, the function should return an empty list.

For instance, in this code sample,

```python
   count_word_lengths('the quick  brown   fox jumped over     the lazy  dog') == [3, 5, 5, 3, 6, 4, 3, 4, 3]`
```

the input string consists of nine (9) words whose respective lengths are shown in the list.

<!-- grade_id: count_word_lengths, points: 0, grade: false, solution: true -->
```python
def count_word_lengths(s):
    assert all([x.isalpha() or x == ' ' for x in s])
    assert type(s) is str
    ans=[]
    s=s.split(" ")
    for i in range(len(s)):
        if len(s[i])!=0:
            ans.append(len(s[i]))

    ###
    return ans
    ### YOUR CODE HERE
    ###
qbf_str = 'the quick brown fox jumped over the lazy dog'
count_word_lengths('')
```

Output:
```
[]
```

<!-- grade_id: count_word_lengths_test, points: 2, grade: true, solution: false -->
```python
# `count_word_lengths_test`: Test cell

# Test 1: Example
qbf_str = 'the quick brown fox jumped over the lazy dog'
qbf_lens = count_word_lengths(qbf_str)
print("Test 1: count_word_lengths('{}') == {}".format(qbf_str, qbf_lens))
assert qbf_lens == [3, 5, 5, 3, 6, 4, 3, 4, 3]

# Test 2: Random strings
from random import choice # 3.5.2 does not have `choices()` (available in 3.6+)
#return ''.join([choice('abcdefghijklmnopqrstuvwxyz') for _ in range(n)])

def random_letter_or_space(pr_space=0.15):
    from random import choice, random
    is_space = (random() <= pr_space)
    if is_space:
        return ' '
    return random_letter()

S_LEN = 40
W_SPACE = 1 / 6
rand_str = random_string(S_LEN, fun=random_letter_or_space)
rand_lens = count_word_lengths(rand_str)
print("Test 2: count_word_lengths('{}') == '{}'".format(rand_str, rand_lens))
c = 0
while c < len(rand_str) and rand_str[c] == ' ':
    c += 1
for k in rand_lens:
    print("  => '{}'".format (rand_str[c:c+k]))
    assert (c+k) == len(rand_str) or rand_str[c+k] == ' '
    c += k
    while c < len(rand_str) and rand_str[c] == ' ':
        c += 1
    
# Test 3: Empty string
print("Test 3: Empty strings...")
assert count_word_lengths('') == []
assert count_word_lengths('   ') == []

print("\n(Passed!)")
```

Output:
```
Test 1: count_word_lengths('the quick brown fox jumped over the lazy dog') == [3, 5, 5, 3, 6, 4, 3, 4, 3]
Test 2: count_word_lengths('f z euoovnkt al cjkknyly wmnsqxza xymzys') == '[1, 1, 8, 2, 8, 8, 6]'
  => 'f'
  => 'z'
  => 'euoovnkt'
  => 'al'
  => 'cjkknyly'
  => 'wmnsqxza'
  => 'xymzys'
Test 3: Empty strings...

(Passed!)
```

**Fin!** You've reached the end of this part. Don't forget to restart and run all cells again to make sure it's all working when run in sequence; and make sure your work passes the submission process. Good luck!